# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
"""
    Find the longest string in a list of strings and return its position number.

    Args:
        L (list): A list of strings.

    Returns:
        int: The position number (index) of the longest string. Indexing starts from 1.

    Examples:
        >>> strings = ['happy', 'birthday', 'to', 'me']
        >>> longest(strings)
        1
        """
def longest(L):
    if not L:
        return None  # Handle an empty list case

    longest_index = 0  # Initialize the index of the longest string
    max_length = len(L[0])  # Initialize the length of the longest string

    for i, word in enumerate(L):
        if len(word) > max_length:
            max_length = len(word)
            longest_index = i

    return longest_index


### END SOLUTION
```

```python
longest(["happy", "birthday", "to", "me"])
```

Output:
```
1
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c(temp_f):
  """(float) -> float
   Converts a temperature in degrees Fahrenheit to degrees Celsius.

  >>>f_to_c(212)
  100.0

  >>>f_to_c(32)
  0.0

  >>>f_to_c(98.6)
  37.0
  """
  return (5/9 * (temp_f - 32))
### END SOLUTION
```

```python
### BEGIN SOLUTION
def f_to_c(tempf):
    """
    Convert temperature from degrees Fahrenheit to degrees Celsius.

    Args:
        tempf (float): Temperature in degrees Fahrenheit.

    Returns:
        float: Temperature in degrees Celsius.
    """
    tempc = 5/9 * (tempf - 32)
    return tempc


 ### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

Output:
```
100.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION


def los(admit_date, discharge_date, diagnosis):
  """Calculates the length of stay in days between the admit and discharge dates.


   # admit_date: The date the patient was admitted to the hospital.
   # discharge_date: The date the patient was discharged from the hospital.
   # diagnosis: The patient's diagnosis.

  Returns:
    The length of stay in days, or 0 if the diagnosis is "Observation".
  """

  if diagnosis == "Observation":
    return 0

  delta = discharge_date - admit_date
  return delta.days



### END SOLUTION
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python

### BEGIN SOLUTION
def average_los(encounters):
    """
    Calculate the average length of stay for encounters in the provided dictionary.

    Args:
        encounters (dict): A dictionary where each key is an encounter ID and the value is a dictionary
            containing at least "admit" (admit date), "diagnosis", and "discharge" (discharge date).

    Returns:
        float: The average length of stay in days, rounded to two decimal places. Excludes encounters with a LOS of 0.
    """
    total_los = 0
    count = 0

    for encounter_id, data in encounters.items():
        admit_date = data.get("admit")
        discharge_date = data.get("discharge")
        diagnosis = data.get("diagnosis")

        if admit_date and discharge_date and diagnosis != "Observation":
            los = (discharge_date - admit_date).days
            if los > 0:
                total_los += los
                count += 1

    if count == 0:
        return 0  # Avoid division by zero

    average = total_los / count
    return round(average, 2)



### END SOLUTION
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
round(average_los(encounters),2)
```

Output:
```
3.33
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(degree_celsius):
    """(int(celsius)) -> (float(farenheit))
    This function coverts given inputs in degree celsius to farenheit.

    >>> c_to_f(100)
    212.0

    >>> c_to_f(0)
    32.0

    >>> c_to_f(37)
    98.60000000000001

    """
    ## calculating farenheit degrees, if the celcius is given

    degree_farenheit = (9/5*degree_celsius) + 32
    return degree_farenheit
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.60000000000001
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.